---
  title: "example_python_challenges"
  output: "html_document"
  v2: false

---
## Choose the right answer

```yaml
type: MultipleChoiceChallenge 
xp: 0 
key: 115qFYuQ4pPlifqZze0QkDORLwVufanGIfC   
```















---
## concatenating lists

```yaml
type: BlanksChallenge 
xp: 0 
key: 8a565a9166   
```















---
## popping lists (2)

```yaml
type: OutputChallenge 
xp: 0 
key: 6258197972   
```















---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 544968a5af   
```

qdsadasdsa

`@instructions`
das
afasfas
adcf

`@hint`
234223

`@pre_exercise_code`

```{undefined}
asd
```


`@sct`

```{undefined}
sdsd
```






---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 4a74954ec6   
```

xvzdgzdfz


`@hint`
zg





`@possible_answers`
dfvzgdf

`@feedbacks`
dfdgfz



